# Excercise 2: Matrices

Consider the Vandermonde matrix given by:

$V =  \left[ \begin{array}{cccc}
1 & x_1 & x_1^2 & x_1^3  \\
1 & x_2 & x_2^2 & x_2^3 \\
1 & x_3 & x_3^2 & x_3^3 \end{array} \right]$

Code has already been provided to define a function named Vandermonde that accepts an N-element vector and assigns it to the variable x. Add commands to generate the Vandermonde matrix from this vector and assign it to the output variable V. For example, if x = [2 3 4], then the resulting matrix, V would be:

$V =  \left[ \begin{array}{cccc}
1 & 2 & 4 & 8  \\
1 & 3 & 9 & 27 \\
1 & 4 & 16 & 64 \end{array} \right]$

Solve this problem using vectorized code (i.e. do not use a loop in your solution.)  Note the variable x is defined as a function input.  Do not overwrite the value in your code.

<b>The mark of this excercise does not count toward your final grade, but submission is compulsory.</b>

There are 4 autograded tests, you can see all of them (i.e. no hidden tests). You can submit unlimited number of times. After each submission, you will be able to see marks and feedback in the Grading Report (View->Grading Report).

Any changes that you make outside of the designated lines will be overwritten on submission.

```python
import numpy as np
```

```python
def Vandermonde(numbers):
    ###
    ### YOUR CODE HERE
    ###
    return v
Vandermonde([2,3,4])#call your function here
```

<!-- grade_id: vandvalues, points: 2, grade: true, solution: false -->
```python
#Test 1 and 2 - checking the result
np.testing.assert_array_equal(Vandermonde([4,5,6]),np.vander([4,5,6], N=4, increasing=True))
np.testing.assert_array_equal(Vandermonde([0,-1,3,5]),np.vander([0,-1,3,5], N=4, increasing=True))
```

<!-- grade_id: np.vand, points: 1, grade: true, solution: false -->
```python
#Test 3 - check that inbuilt numpy.vander is not used
from unittest.mock import patch
with patch('numpy.vander') as inbuiltVand:
    Vandermonde([1,2,3])
inbuiltVand.assert_not_called()#check that inbuilt math factorial not used
```