# Excercise 1: Factorial

In this simple diagnostic excercise you will have to use conditionals, loops and ValueError. Your task is to <b>create a function to calculate a factorial of a number, using a loop</b>. Raise ValueError if the number is non-negative.

<b>The mark of this excercise does not count toward your final grade, but submission is compulsory.</b>

There are 4 autograded tests, you can see all of them (i.e. no hidden tests). You can submit unlimited number of times. After each submission, you will be able to see marks and feedback in the Grading Report (View->Grading Report).

Any changes that you make outside of the designated lines will be overwritten on submission.

```python
def factorial(number):
    ###
    ### YOUR CODE HERE
    if (number < 0):
        raise ValueError
    result = 1
    for i in range(1, number + 1):
        result *= i
    ###
    return result
```

<!-- grade_id: values, points: 2, grade: true, solution: false -->
```python
#Test 1 and 2 - checking the result
from nose.tools import assert_equal
from math import factorial as mfact 
assert_equal(factorial(2),2)
assert_equal(factorial(10),mfact(10))
```

<!-- grade_id: exceptions, points: 1, grade: true, solution: false -->
```python
#Test3 - check presence of error assertion
from nose.tools import assert_raises
assert_raises(ValueError,factorial,-2)
```

<!-- grade_id: math_factorial, points: 1, grade: true, solution: false -->
```python
#Test 4 - check that inbuilt math.factorial is not used
from unittest.mock import patch
with patch('math.factorial') as mock_factorial:
    factorial(3)
mock_factorial.assert_not_called()#check that inbuilt math factorial not used
```